<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="coderunner">
    <name>
      <text>Session 4 Iris Functions Submission</text>
    </name>

    <questiontext format="html">
      <text><![CDATA[
<p dir="ltr" style="text-align: left;">
Submit the complete Session 4 code for refactoring the Iris classifier into functions.
</p>

<p>Start from the provided Session 4 template and complete the required functions exactly as described in <code>session4/ses4.tex</code>.</p>

<p>The grader will run 10 checks across these functions and the end-to-end program flow:</p>
<pre>make_print_status
setup_application_list
compute_threshold_prediction
derive_true_label
update_result_counts
calculate_accuracy
run_prediction_loop
print_summary
main</pre>

<p>Keep the function names, return values, and printed output format consistent with the tutorial. Your final program should still produce the Session 4 status lines, prediction loop trace, summary, and the final end marker.</p>
      ]]></text>
    </questiontext>

    <generalfeedback format="html">
      <text><![CDATA[
<p>Common mistakes for this submission include:</p>
<ul>
  <li>Leaving <code>pass</code> in <code>make_print_status</code>, <code>compute_threshold_prediction</code>, or <code>update_result_counts</code></li>
  <li>Forgetting to define <code>flower2</code> or forgetting to include it in the dataset list</li>
  <li>Missing <code>y_true = derive_true_label(sample)</code> inside <code>run_prediction_loop</code></li>
  <li>Returning the wrong values or the wrong order from <code>run_prediction_loop</code></li>
  <li>Forgetting <code>return accuracy</code> in the final <code>calculate_accuracy(correct, total)</code> function</li>
  <li>Changing the required print format for the per-sample trace or summary lines</li>
  <li>Using the wrong status text instead of <code>Build dataset</code>, <code>Run prediction loop</code>, or <code>Print summary</code></li>
</ul>
      ]]></text>
    </generalfeedback>

    <defaultgrade>10</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>

    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>0</allornothing>
    <penaltyregime>10, 20, ...</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>

    <answerboxlines>85</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>

    <answerpreload><![CDATA["""Session 4: Move Session 3 code into functions.
rpb
This version keeps the same rule from Session 3.
The goal is to help a beginner see where each part came from.
We keep the code simple and close to the original step-by-step style.
"""
LABEL_KEY = "species"
THRESHOLD = 2.0
FEATURE_NAME = "petal_length"
POSITIVE_LABEL = "setosa"
NEGATIVE_LABEL = "not_setosa"


# Task 1: Create a helper function to print status updates
def make_print_status(status_text):
    """Print a small status message.

    Args:
        status_text (str): A short message to show what the program is doing.
    """
    pass

# Task 2: Create the flower dataset


def setup_application_list():
    """Combination of Task I and Task II in session III, but now in a function."""

    # Task I in session III: Define dictionaries for flower1 and flower2 using canonical keys
    # Paste the two dictionaries you created in Session 3 here, but make sure to use the same keys as in the original dataset (id, sepal_length, sepal_width, petal_length, petal_width, species). Here, we already define the flower1, but you should also define flower2 with the same keys.
    flower1 = {
        "id": "flower1",
        "sepal_length": 5.1,
        "sepal_width": 3.5,
        "petal_length": 1.4,
        "petal_width": 0.2,
        "species": "setosa"
    }

    # flower2 = {
    #     "id": "flower2",
    # }

    # Task 2 in session 3: Build the dataset list
    # Combine our dictionaries into a single list
    dataset = [flower1
               #    , flower2
               ]
    print("Dataset:", dataset)
    # Note here that we return the dataset list from this function, so we can use it later in the main function.
    return dataset

# Task 5: Predict the class using petal length


def compute_threshold_prediction(sample):
    """Predict the label for one flower sample.

    This uses the same rule from Session 3:
    if petal_length is less than 2.0, predict "setosa".
    Otherwise, predict "not_setosa".

    Args:
        sample (dict): One flower sample.

    Returns:
        str: The predicted label.
    """
    # we deliberately put the pass here, and student need to replace it with the if statement and return statement to make the prediction.
    # This should be given in the instruction, on what they should paste. sub task 5, the need to also give the return statement, which is missing in the instruction now.
    pass

# Task 6: Convert species into the lesson’s binary label


def derive_true_label(sample):
    """Convert the real species into the lesson label.

    In this lesson, we only use two labels:
    - "setosa"
    - "not_setosa"

    Args:
        sample (dict): One flower sample.

    Returns:
        str: The true label for this lesson.
    """

    if sample[LABEL_KEY] == POSITIVE_LABEL:
        y_true = POSITIVE_LABEL
    else:
        y_true = NEGATIVE_LABEL

    return y_true

# Task 7: Update prediction counts and save results


def update_result_counts(correct, wrong, total, y_pred_list, y_pred, y_true):
    """Update the counters and prediction list for one sample.

    Args:
        correct (int): Current number of correct predictions.
        wrong (int): Current number of wrong predictions.
        total (int): Current number of processed samples.
        y_pred_list (list): Current list of predictions.
        y_pred (str): Predicted label.
        y_true (str): True label.

    Returns:
        tuple: Updated correct, wrong, total, y_pred_list
    """
    # we will provide only a pass, but we drop all of below code, and student need to write the if statement to compare y_pred and y_true, and update the correct and wrong counts accordingly. They also need to remember to update the total count and append the y_pred to the y_pred_list.
    pass

# Task 10: Compute the overall accuracy


def calculate_accuracy(correct, total):
    """Calculate accuracy percentage.

        Args:
        correct (int): Number of correct predictions.
        total (int): Number of processed samples.

        Returns:
                float: Accuracy percentage.
    """
    if total > 0:
        accuracy = (correct / total) * 100
    else:
        accuracy = 0.0

    # They need to provide the return value, maybe, we can ask them, what is missing in this function, go and check the function call in the main function, and see what they need to return here to make the main function work.


# Entry to Task 4 until Task 9: Run the same prediction loop from Session 3, Wrap the prediction workflow in a function


def run_prediction_loop(dataset):
    """Run the same prediction loop from Session 3, but inside a function.

    This function keeps the beginner-friendly variables:
    correct, wrong, total, and y_pred_list.

    Args:
        dataset (list): A list of flower dictionaries.

    Returns:
        tuple: correct, wrong, total, y_pred_list, accuracy
    """
    # Task III from Session III: initialize metrics and predictions
    correct = 0      # Count of correct predictions
    wrong = 0        # Count of wrong predictions
    total = 0        # Total samples processed
    y_pred_list = []  # List of all predictions made

    print("\n=== Start session 4 Prediction Loop ===")

    # Task IV and Task V from Session III
    for sample in dataset:
        print("\nProcessing sample with id:", sample["id"])

        # Task 5: Predict the class using petal length
        y_pred = compute_threshold_prediction(sample)

        # Task 6: Convert species into the lesson’s binary label
        # For task 6, student are require to create the function call and function definition:y_true = derive_true_label(sample)

        # Task 7: Update prediction counts and save results
        # For task 7, we will not provide the correct, wrong, total, y_pred_list = update_result_counts( correct, wrong, total, y_pred_list, y_pred, y_true ), even the function body is only a pass
        correct, wrong, total, y_pred_list = update_result_counts(
            correct, wrong, total, y_pred_list, y_pred, y_true)

        # Task 8: Display the result for each sample
        # The just need to uncomment the print statement below to see the result for each sample,
        # print(
        #     f"id={sample['id']} | true={y_true} | pred={y_pred} | "
        #     f"petal_length={sample['petal_length']}"
        # )
    # Task 9: Return the prediction loop results
    return < your_code > , wrong, < your_code > , < your_code >

# Task 10: Calculate accuracy


def calculate_accuracy(correct, total):
    """Calculate accuracy percentage.

    Args:
        correct (int): Number of correct predictions.
        total (int): Number of processed samples.

    Returns:
        float: Accuracy percentage.
    """
    if total > 0:
        accuracy = (correct / total) * 100
    else:
        accuracy = 0.0


# Task 12: Create and call the summary-report function


def print_summary(correct, wrong, total, y_pred_list, accuracy):
    """Print the final results after the loop is finished.

    Args:
        correct (int): Number of correct predictions.
        wrong (int): Number of wrong predictions.
        total (int): Number of processed samples.
        y_pred_list (list): List of predicted labels.
        accuracy (float): Accuracy percentage.
    """
    print("\n=== session 4 Summary ===")
    print("Correct:", correct)
    print("Wrong:", wrong)
    print("Total:", total)
    print("Accuracy (%):", round(accuracy, 2))
    print("All predictions:", y_pred_list)


def main():
    """Run the full beginner version of the program."""

    # Task 1 : Create a helper function to print status updates
    make_print_status("Build dataset")

    # Task 2: Create the flower dataset
    dataset = setup_application_list()

    # Task 3 : Show a status update before running prediction
    # Uncomment this line to see status messages
    make_print_status("Run prediction loop")

    # Entry to Task 4 until Task 9: Wrap the prediction workflow in a function
    # correct, <your_code>, total, y_pred_list = run_prediction_loop(dataset)
    correct, wrong, total, y_pred_list = run_prediction_loop(dataset)

    # Task 10: Calculate accuracy
    accuracy = calculate_accuracy(correct, total)

    # Task 11: Show a status update before printing the summary
    make_print_status( < your_code > )

    # Task 12: Create and call the summary-report function
    print_summary(correct, wrong, total, y_pred_list, accuracy)


if __name__ == "__main__":
    main()
]]></answerpreload>

    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template><![CDATA[import contextlib
import io
import json

student_code = """{{ STUDENT_ANSWER | e('py') }}"""
rows = [[
    "Check",
    "Expected",
    "Got",
    "Comment",
    "Mark",
    "iscorrect"
]]
total_awarded = 0.0


def compact(value):
    text = str(value)
    if len(text) > 220:
        return text[:217] + "..."
    return text


def cleaned_lines(text):
    return [line.rstrip() for line in text.splitlines() if line.strip()]


def add_result(label, expected, got, comment, awarded, mark, iscorrect):
    global total_awarded
    total_awarded += awarded
    rows.append([
        label,
        compact(expected),
        compact(got),
        comment,
        f"{awarded:.1f}/{mark:.1f}",
        iscorrect
    ])


def require_callable(namespace, name):
    obj = namespace.get(name)
    if not callable(obj):
        raise NameError(f"Function '{name}' is missing or not callable.")
    return obj


expected_dataset = [
    {
        "id": "flower1",
        "sepal_length": 5.1,
        "sepal_width": 3.5,
        "petal_length": 1.4,
        "petal_width": 0.2,
        "species": "setosa",
    },
    {
        "id": "flower2",
        "sepal_length": 4.9,
        "sepal_width": 3.0,
        "petal_length": 1.4,
        "petal_width": 0.2,
        "species": "setosa",
    },
]


def run_test(label, mark, expected, check_fn):
    try:
        success, got, comment = check_fn()
        awarded = mark if success else 0.0
        add_result(label, expected, got, comment, awarded, mark, 1 if success else 0)
    except Exception as exc:
        add_result(
            label,
            expected,
            f"{type(exc).__name__}: {exc}",
            "This check raised an exception.",
            0.0,
            mark,
            0
        )


try:
    compiled = compile(student_code, "<student>", "exec")
    namespace = {"__name__": "student_submission"}
    with contextlib.redirect_stdout(io.StringIO()):
        exec(compiled, namespace)

    def check_make_print_status():
        func = require_callable(namespace, "make_print_status")
        stdout_buffer = io.StringIO()
        with contextlib.redirect_stdout(stdout_buffer):
            func("Build dataset")
        lines = cleaned_lines(stdout_buffer.getvalue())
        expected_lines = ["[STATUS] Build dataset"]
        success = lines == expected_lines
        return success, lines, "Expected one exact status line."

    def check_setup_application_list():
        func = require_callable(namespace, "setup_application_list")
        stdout_buffer = io.StringIO()
        with contextlib.redirect_stdout(stdout_buffer):
            dataset = func()
        lines = cleaned_lines(stdout_buffer.getvalue())
        success = (
            dataset == expected_dataset
            and len(lines) == 1
            and lines[0].startswith("Dataset:")
            and "flower1" in lines[0]
            and "flower2" in lines[0]
        )
        got = {"returned": dataset, "printed": lines}
        return success, got, "Function should return both flowers and print the dataset line."

    def check_compute_prediction_setosa():
        func = require_callable(namespace, "compute_threshold_prediction")
        sample = {"petal_length": 1.4}
        result = func(sample)
        success = result == "setosa"
        return success, result, "Short petals should predict 'setosa'."

    def check_compute_prediction_not_setosa():
        func = require_callable(namespace, "compute_threshold_prediction")
        sample = {"petal_length": 3.5}
        result = func(sample)
        success = result == "not_setosa"
        return success, result, "Longer petals should predict 'not_setosa'."

    def check_derive_true_label():
        func = require_callable(namespace, "derive_true_label")
        setosa_result = func({"species": "setosa"})
        other_result = func({"species": "versicolor"})
        success = setosa_result == "setosa" and other_result == "not_setosa"
        got = {"setosa": setosa_result, "versicolor": other_result}
        return success, got, "True labels should map to the lesson's binary labels."

    def check_update_result_counts():
        func = require_callable(namespace, "update_result_counts")
        first = func(0, 0, 0, [], "setosa", "setosa")
        second = func(1, 0, 1, ["setosa"], "not_setosa", "setosa")
        success = first == (1, 0, 1, ["setosa"]) and second == (1, 1, 2, ["setosa", "not_setosa"])
        got = {"correct_case": first, "wrong_case": second}
        return success, got, "Counters, total, and prediction list should update correctly."

    def check_calculate_accuracy():
        func = require_callable(namespace, "calculate_accuracy")
        normal = func(2, 2)
        zero_total = func(0, 0)
        success = normal == 100.0 and zero_total == 0.0
        got = {"two_of_two": normal, "zero_total": zero_total}
        return success, got, "Accuracy should be 100.0 for 2/2 and 0.0 when total is zero."

    def check_run_prediction_loop():
        func = require_callable(namespace, "run_prediction_loop")
        stdout_buffer = io.StringIO()
        with contextlib.redirect_stdout(stdout_buffer):
            result = func(expected_dataset)
        lines = cleaned_lines(stdout_buffer.getvalue())
        required_lines = [
            "=== Start session 4 Prediction Loop ===",
            "Processing sample with id: flower1",
            "id=flower1 | true=setosa | pred=setosa | petal_length=1.4",
            "Processing sample with id: flower2",
            "id=flower2 | true=setosa | pred=setosa | petal_length=1.4",
        ]
        success = result == (2, 0, 2, ["setosa", "setosa"]) and all(line in lines for line in required_lines)
        got = {"returned": result, "printed": lines}
        return success, got, "Loop should return counts and print both trace lines."

    def check_print_summary():
        func = require_callable(namespace, "print_summary")
        stdout_buffer = io.StringIO()
        with contextlib.redirect_stdout(stdout_buffer):
            func(2, 0, 2, ["setosa", "setosa"], 100.0)
        lines = cleaned_lines(stdout_buffer.getvalue())
        expected_lines = [
            "=== session 4 Summary ===",
            "Correct: 2",
            "Wrong: 0",
            "Total: 2",
            "Accuracy (%): 100.0",
            "All predictions: ['setosa', 'setosa']",
        ]
        success = lines == expected_lines
        return success, lines, "Summary should print the six expected lines."

    def check_main():
        func = require_callable(namespace, "main")
        stdout_buffer = io.StringIO()
        with contextlib.redirect_stdout(stdout_buffer):
            func()
        lines = cleaned_lines(stdout_buffer.getvalue())
        required_lines = [
            "[STATUS] Build dataset",
            "[STATUS] Run prediction loop",
            "=== Start session 4 Prediction Loop ===",
            "Processing sample with id: flower1",
            "id=flower1 | true=setosa | pred=setosa | petal_length=1.4",
            "Processing sample with id: flower2",
            "id=flower2 | true=setosa | pred=setosa | petal_length=1.4",
            "[STATUS] Print summary",
            "=== session 4 Summary ===",
            "Correct: 2",
            "Wrong: 0",
            "Total: 2",
            "Accuracy (%): 100.0",
            "All predictions: ['setosa', 'setosa']",
            "=== End of session 4 ===",
        ]
        success = all(line in lines for line in required_lines)
        return success, lines, "main() should run the full Session 4 workflow."

    run_test(
        "1. make_print_status",
        1.0,
        "Prints [STATUS] Build dataset",
        check_make_print_status
    )
    run_test(
        "2. setup_application_list",
        1.0,
        "Returns the two-flower dataset and prints Dataset: ...",
        check_setup_application_list
    )
    run_test(
        "3. compute_threshold_prediction short petal",
        1.0,
        "Returns setosa for petal_length < 2.0",
        check_compute_prediction_setosa
    )
    run_test(
        "4. compute_threshold_prediction long petal",
        1.0,
        "Returns not_setosa for petal_length >= 2.0",
        check_compute_prediction_not_setosa
    )
    run_test(
        "5. derive_true_label",
        1.0,
        "Maps setosa -> setosa and non-setosa -> not_setosa",
        check_derive_true_label
    )
    run_test(
        "6. update_result_counts",
        1.0,
        "Updates counters, total, and y_pred_list",
        check_update_result_counts
    )
    run_test(
        "7. calculate_accuracy",
        1.0,
        "Returns 100.0 for 2/2 and 0.0 for zero total",
        check_calculate_accuracy
    )
    run_test(
        "8. run_prediction_loop",
        1.0,
        "Returns (2, 0, 2, ['setosa', 'setosa']) and prints both trace lines",
        check_run_prediction_loop
    )
    run_test(
        "9. print_summary",
        1.0,
        "Prints the six summary lines",
        check_print_summary
    )
    run_test(
        "10. main",
        1.0,
        "Runs the full Session 4 program and prints the required workflow lines",
        check_main
    )

    fraction = total_awarded / 10.0
except Exception as exc:
    rows.append([
        "Program execution",
        "Student code compiles and loads successfully",
        f"{type(exc).__name__}: {exc}",
        "The code could not be imported for function-level testing.",
        "0.0/10.0",
        0
    ])
    fraction = 0.0

print(json.dumps({
    "fraction": fraction,
    "testresults": rows
}))]]></template>
    <iscombinatortemplate>1</iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA["""Session 4: Move Session 3 code into functions.

This version keeps the same rule from Session 3.
The goal is to help a beginner see where each part came from.
We keep the code simple and close to the original step-by-step style.
"""
LABEL_KEY = "species"
THRESHOLD = 2.0
FEATURE_NAME = "petal_length"
POSITIVE_LABEL = "setosa"
NEGATIVE_LABEL = "not_setosa"


# Task 1: Create a helper function to print status updates
def make_print_status(status_text):
    """Print a small status message.

    Args:
        status_text (str): A short message to show what the program is doing.
    """
    # student should paste the statement print(f"[STATUS] {status_text}")
    print(f"[STATUS] {status_text}")
    # Comment out the print statement for now, but you can uncomment it to see status messages.

# Task 2: Create the flower dataset


def setup_application_list():
    """Combination of Task I and Task II in session III, but now in a function."""

    # Task I in session III: Define dictionaries for flower1 and flower2 using canonical keys
    # Paste the two dictionaries you created in Session 3 here, but make sure to use the same keys as in the original dataset (id, sepal_length, sepal_width, petal_length, petal_width, species). Here, we already define the flower1, but you should also define flower2 with the same keys.
    flower1 = {
        "id": "flower1",
        "sepal_length": 5.1,
        "sepal_width": 3.5,
        "petal_length": 1.4,
        "petal_width": 0.2,
        "species": "setosa"
    }

    flower2 = {
        "id": "flower2",
        "sepal_length": 4.9,
        "sepal_width": 3.0,
        "petal_length": 1.4,
        "petal_width": 0.2,
        "species": "setosa"
    }

    # Task 2 in session 3: Build the dataset list
    # Combine our dictionaries into a single list
    dataset = [flower1, flower2]
    print("Dataset:", dataset)
    # Note here that we return the dataset list from this function, so we can use it later in the main function.
    return dataset

# Task 5: Predict the class using petal length


def compute_threshold_prediction(sample):
    """Predict the label for one flower sample.

    This uses the same rule from Session 3:
    if petal_length is less than 2.0, predict "setosa".
    Otherwise, predict "not_setosa".

    Args:
        sample (dict): One flower sample.

    Returns:
        str: The predicted label.
    """
    # we deliberately put the pass here, and student need to replace it with the if statement and return statement to make the prediction.
    # This should be given in the instruction, on what they should paste. sub task 5, the need to also give the return statement, which is missing in the instruction now.
    if sample[FEATURE_NAME] < THRESHOLD:
        y_pred = POSITIVE_LABEL
    else:
        y_pred = NEGATIVE_LABEL

    return y_pred

# Task 6: Convert species into the lesson’s binary label


def derive_true_label(sample):
    """Convert the real species into the lesson label.

    In this lesson, we only use two labels:
    - "setosa"
    - "not_setosa"

    Args:
        sample (dict): One flower sample.

    Returns:
        str: The true label for this lesson.
    """

    if sample[LABEL_KEY] == POSITIVE_LABEL:
        y_true = POSITIVE_LABEL
    else:
        y_true = NEGATIVE_LABEL

    return y_true

# Task 7: Update prediction counts and save results


def update_result_counts(correct, wrong, total, y_pred_list, y_pred, y_true):
    """Update the counters and prediction list for one sample.

    Args:
        correct (int): Current number of correct predictions.
        wrong (int): Current number of wrong predictions.
        total (int): Current number of processed samples.
        y_pred_list (list): Current list of predictions.
        y_pred (str): Predicted label.
        y_true (str): True label.

    Returns:
        tuple: Updated correct, wrong, total, y_pred_list
    """
    # we will provide only a pass, but we drop all of below code, and student need to write the if statement to compare y_pred and y_true, and update the correct and wrong counts accordingly. They also need to remember to update the total count and append the y_pred to the y_pred_list.
    if y_pred == y_true:
        correct += 1
    else:
        wrong += 1

    total += 1
    y_pred_list.append(y_pred)

    return correct, wrong, total, y_pred_list

# Task 10: Compute the overall accuracy


def calculate_accuracy(correct, total):
    """Calculate accuracy percentage.

        Args:
        correct (int): Number of correct predictions.
        total (int): Number of processed samples.

        Returns:
                float: Accuracy percentage.
    """
    if total > 0:
        accuracy = (correct / total) * 100
    else:
        accuracy = 0.0

    # They need to provide the return value, maybe, we can ask them, what is missing in this function, go and check the function call in the main function, and see what they need to return here to make the main function work.
    return accuracy

# Entry to Task 4 until Task 9: Run the same prediction loop from Session 3, Wrap the prediction workflow in a function


def run_prediction_loop(dataset):
    """Run the same prediction loop from Session 3, but inside a function.

    This function keeps the beginner-friendly variables:
    correct, wrong, total, and y_pred_list.

    Args:
        dataset (list): A list of flower dictionaries.

    Returns:
        tuple: correct, wrong, total, y_pred_list, accuracy
    """
    # Task III from Session III: initialize metrics and predictions
    correct = 0      # Count of correct predictions
    wrong = 0        # Count of wrong predictions
    total = 0        # Total samples processed
    y_pred_list = []  # List of all predictions made

    print("\n=== Start session 4 Prediction Loop ===")

    # Task IV and Task V from Session III
    for sample in dataset:
        print("\nProcessing sample with id:", sample["id"])

        # Task 5: Predict the class using petal length
        y_pred = compute_threshold_prediction(sample)

        # Task 6: Convert species into the lesson’s binary label
        # For task 6, we purposely mention, they need to write do the y_true = derive_true_label(sample), but the function body of the derive_true_label is provided as it is.
        y_true = derive_true_label(sample)

        # Task 7: Update prediction counts and save results
        # For task 7, we will not provide the correct, wrong, total, y_pred_list = update_result_counts( correct, wrong, total, y_pred_list, y_pred, y_true ), even the function body is only a pass
        correct, wrong, total, y_pred_list = update_result_counts(
            correct, wrong, total, y_pred_list, y_pred, y_true
        )

        # Task 8: Display the result for each sample
        # The just need to uncomment the print statement below to see the result for each sample,
        print(
            f"id={sample['id']} | true={y_true} | pred={y_pred} | "
            f"petal_length={sample['petal_length']}"
        )
    # Task 9: Return the prediction loop results
    # student need to file <your_code>,wrong, <your_code>,<your_code>
    return correct, wrong, total, y_pred_list

# Task 10: Calculate accuracy


def calculate_accuracy(correct, total):
    """Calculate accuracy percentage.

    Args:
        correct (int): Number of correct predictions.
        total (int): Number of processed samples.

    Returns:
        float: Accuracy percentage.
    """
    if total > 0:
        accuracy = (correct / total) * 100
    else:
        accuracy = 0.0
    # Student need to provide the return value, maybe, we can ask them, what is missing in this function, go and check the function call in the main function, and see what they need to return here to make the main function work.
    return accuracy

# Task 12: Create and call the summary-report function


def print_summary(correct, wrong, total, y_pred_list, accuracy):
    """Print the final results after the loop is finished.

    Args:
        correct (int): Number of correct predictions.
        wrong (int): Number of wrong predictions.
        total (int): Number of processed samples.
        y_pred_list (list): List of predicted labels.
        accuracy (float): Accuracy percentage.
    """
    print("\n=== session 4 Summary ===")
    print("Correct:", correct)
    print("Wrong:", wrong)
    print("Total:", total)
    print("Accuracy (%):", round(accuracy, 2))
    print("All predictions:", y_pred_list)


def main():
    """Run the full beginner version of the program."""

    # Task 1 : Create a helper function to print status updates
    make_print_status("Build dataset")

    # Task 2: Create the flower dataset
    dataset = setup_application_list()

    # Task 3 : Show a status update before running prediction
    # Uncomment this line to see status messages
    make_print_status("Run prediction loop")

    # Entry to Task 4 until Task 9: Wrap the prediction workflow in a function
    # correct, <your_code>, total, y_pred_list = run_prediction_loop(dataset)
    correct, wrong, total, y_pred_list = run_prediction_loop(dataset)

    # Task 10: Calculate accuracy
    accuracy = calculate_accuracy(correct, total)

    # Task 11: Show a status update before printing the summary
    # we ask student to write "Print summary" as the status text, replacing the <your_code>.
    make_print_status("Print summary")

    # Task 12: Create and call the summary-report function
    print_summary(correct, wrong, total, y_pred_list, accuracy)
    print("\n=== End of session 4 ===")


if __name__ == "__main__":
    main()
]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader>TemplateGrader</grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>

    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="0.0000000">
        <testcode>
          <text>Run the Session 4 function-based grader with 10 internal checks.</text>
        </testcode>
        <stdin>
          <text></text>
        </stdin>
        <expected>
          <text></text>
        </expected>
        <extra>
          <text></text>
        </extra>
        <display>
          <text>SHOW</text>
        </display>
      </testcase>
    </testcases>

    <tags>
      <tag>
        <text>session4</text>
      </tag>
      <tag>
        <text>session4_submission_iris_functions</text>
      </tag>
      <tag>
        <text>submission</text>
      </tag>
      <tag>
        <text>function-check</text>
      </tag>
      <tag>
        <text>unit-tests</text>
      </tag>
    </tags>
  </question>
</quiz>
