<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="coderunner">
    <name>
      <text>Session 4 Task 7 Iris</text>
    </name>

    <questiontext format="html">
      <text><![CDATA[
<p dir="ltr" style="text-align: left;">Complete <code>update_result_counts(...)</code>.</p><p>The function must update <code>correct</code> or <code>wrong</code>, increase <code>total</code>, append the prediction to <code>y_pred_list</code>, and return all four updated values.</p><p>Paste your current Session 4 file into the answer box. This quiz checks only Task 7.</p>
      ]]></text>
    </questiontext>

    <generalfeedback format="html">
      <text><![CDATA[
<p>Common mistakes for Task 7:</p><ul><li>Forgetting to increase <code>total</code> for every sample</li><li>Updating the wrong counter</li><li>Forgetting to append <code>y_pred</code> to <code>y_pred_list</code></li><li>Returning the values in the wrong order</li></ul>
      ]]></text>
    </generalfeedback>

    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>

    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>0</allornothing>
    <penaltyregime>10, 20, ...</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>

    <answerboxlines>85</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>

    <answerpreload><![CDATA[# Task 7: Update prediction counts and save results
def update_result_counts(correct, wrong, total, y_pred_list, y_pred, y_true):
    """Update the counters and prediction list for one sample.

    Args:
        correct (int): Current number of correct predictions.
        wrong (int): Current number of wrong predictions.
        total (int): Current number of processed samples.
        y_pred_list (list): Current list of predictions.
        y_pred (str): Predicted label.
        y_true (str): True label.

    Returns:
        tuple: Updated correct, wrong, total, y_pred_list
    """
    # student should compare y_pred and y_true,
    # update the counters, increase total,
    # append y_pred, and return all four updated values.
    pass
]]></answerpreload>

    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template><![CDATA[import contextlib
import io
import json
import re

student_code = """{{ STUDENT_ANSWER | e('py') }}"""
rows = [["Check", "Expected", "Got", "Comment", "Mark", "iscorrect"]]
total_awarded = 0.0

expected_dataset = [
    {
        "id": "flower1",
        "sepal_length": 5.1,
        "sepal_width": 3.5,
        "petal_length": 1.4,
        "petal_width": 0.2,
        "species": "setosa",
    },
    {
        "id": "flower2",
        "sepal_length": 4.9,
        "sepal_width": 3.0,
        "petal_length": 1.4,
        "petal_width": 0.2,
        "species": "setosa",
    },
]


def compact(value):
    text = str(value)
    if len(text) > 220:
        return text[:217] + "..."
    return text


def cleaned_lines(text):
    return [line.rstrip() for line in text.splitlines() if line.strip()]


def add_result(label, expected, got, comment, awarded, mark, iscorrect):
    global total_awarded
    total_awarded += awarded
    rows.append([
        label,
        compact(expected),
        compact(got),
        comment,
        f"{awarded:.1f}/{mark:.1f}",
        iscorrect,
    ])


def require_callable(namespace, name):
    obj = namespace.get(name)
    if not callable(obj):
        raise NameError(f"Function '{name}' is missing or not callable.")
    return obj


def sanitize(code, fix_return=True, fix_main_status=True):
    if fix_return:
        code = re.sub(
            r"return\s*<\s*your_code\s*>\s*,\s*wrong\s*,\s*<\s*your_code\s*>\s*,\s*<\s*your_code\s*>",
            "return correct, wrong, total, y_pred_list",
            code,
        )
    if fix_main_status:
        code = re.sub(
            r"make_print_status\(\s*<\s*your_code\s*>\s*\)",
            'make_print_status("Print summary")',
            code,
        )
    return code


def load_namespace(fix_return=True, fix_main_status=True):
    code = sanitize(student_code, fix_return=fix_return, fix_main_status=fix_main_status)
    compiled = compile(code, "<student>", "exec")
    namespace = {"__name__": "student_submission"}
    with contextlib.redirect_stdout(io.StringIO()):
        exec(compiled, namespace)
    return namespace

DEFAULT_GRADE = 1.0

try:
    namespace = load_namespace()
    func = require_callable(namespace, "update_result_counts")
    correct_case = func(0, 0, 0, [], "setosa", "setosa")
    wrong_case = func(1, 0, 1, ["setosa"], "not_setosa", "setosa")
    success = correct_case == (1, 0, 1, ["setosa"]) and wrong_case == (1, 1, 2, ["setosa", "not_setosa"])
    add_result("Task 7", "update_result_counts returns updated counters, total, and prediction list", {"correct_case": correct_case, "wrong_case": wrong_case}, "The function should handle both correct and wrong predictions.", 1.0 if success else 0.0, 1.0, 1 if success else 0)
except Exception as exc:
    add_result("Task 7", "Metric update function works", f"{type(exc).__name__}: {exc}", "The code could not be imported for Task 7.", 0.0, 1.0, 0)

print(json.dumps({
    "fraction": total_awarded / DEFAULT_GRADE,
    "testresults": rows,
}))]]></template>
    <iscombinatortemplate>1</iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA["""Session 4: Move Session 3 code into functions.

This version keeps the same rule from Session 3.
The goal is to help a beginner see where each part came from.
We keep the code simple and close to the original step-by-step style.
"""
LABEL_KEY = "species"
THRESHOLD = 2.0
FEATURE_NAME = "petal_length"
POSITIVE_LABEL = "setosa"
NEGATIVE_LABEL = "not_setosa"


# Task 1: Create a helper function to print status updates
def make_print_status(status_text):
    """Print a small status message.

    Args:
        status_text (str): A short message to show what the program is doing.
    """
    # student should paste the statement print(f"[STATUS] {status_text}")
    print(f"[STATUS] {status_text}")
    # Comment out the print statement for now, but you can uncomment it to see status messages.

# Task 2: Create the flower dataset


def setup_application_list():
    """Combination of Task I and Task II in session III, but now in a function."""

    # Task I in session III: Define dictionaries for flower1 and flower2 using canonical keys
    # Paste the two dictionaries you created in Session 3 here, but make sure to use the same keys as in the original dataset (id, sepal_length, sepal_width, petal_length, petal_width, species). Here, we already define the flower1, but you should also define flower2 with the same keys.
    flower1 = {
        "id": "flower1",
        "sepal_length": 5.1,
        "sepal_width": 3.5,
        "petal_length": 1.4,
        "petal_width": 0.2,
        "species": "setosa"
    }

    flower2 = {
        "id": "flower2",
        "sepal_length": 4.9,
        "sepal_width": 3.0,
        "petal_length": 1.4,
        "petal_width": 0.2,
        "species": "setosa"
    }

    # Task 2 in session 3: Build the dataset list
    # Combine our dictionaries into a single list
    dataset = [flower1, flower2]
    print("Dataset:", dataset)
    # Note here that we return the dataset list from this function, so we can use it later in the main function.
    return dataset

# Task 5: Predict the class using petal length


def compute_threshold_prediction(sample):
    """Predict the label for one flower sample.

    This uses the same rule from Session 3:
    if petal_length is less than 2.0, predict "setosa".
    Otherwise, predict "not_setosa".

    Args:
        sample (dict): One flower sample.

    Returns:
        str: The predicted label.
    """
    # we deliberately put the pass here, and student need to replace it with the if statement and return statement to make the prediction.
    # This should be given in the instruction, on what they should paste. sub task 5, the need to also give the return statement, which is missing in the instruction now.
    if sample[FEATURE_NAME] < THRESHOLD:
        y_pred = POSITIVE_LABEL
    else:
        y_pred = NEGATIVE_LABEL

    return y_pred

# Task 6: Convert species into the lesson’s binary label


def derive_true_label(sample):
    """Convert the real species into the lesson label.

    In this lesson, we only use two labels:
    - "setosa"
    - "not_setosa"

    Args:
        sample (dict): One flower sample.

    Returns:
        str: The true label for this lesson.
    """

    if sample[LABEL_KEY] == POSITIVE_LABEL:
        y_true = POSITIVE_LABEL
    else:
        y_true = NEGATIVE_LABEL

    return y_true

# Task 7: Update prediction counts and save results


def update_result_counts(correct, wrong, total, y_pred_list, y_pred, y_true):
    """Update the counters and prediction list for one sample.

    Args:
        correct (int): Current number of correct predictions.
        wrong (int): Current number of wrong predictions.
        total (int): Current number of processed samples.
        y_pred_list (list): Current list of predictions.
        y_pred (str): Predicted label.
        y_true (str): True label.

    Returns:
        tuple: Updated correct, wrong, total, y_pred_list
    """
    # we will provide only a pass, but we drop all of below code, and student need to write the if statement to compare y_pred and y_true, and update the correct and wrong counts accordingly. They also need to remember to update the total count and append the y_pred to the y_pred_list.
    if y_pred == y_true:
        correct += 1
    else:
        wrong += 1

    total += 1
    y_pred_list.append(y_pred)

    return correct, wrong, total, y_pred_list

# Task 10: Compute the overall accuracy


def calculate_accuracy(correct, total):
    """Calculate accuracy percentage.

        Args:
        correct (int): Number of correct predictions.
        total (int): Number of processed samples.

        Returns:
                float: Accuracy percentage.
    """
    if total > 0:
        accuracy = (correct / total) * 100
    else:
        accuracy = 0.0

    # They need to provide the return value, maybe, we can ask them, what is missing in this function, go and check the function call in the main function, and see what they need to return here to make the main function work.
    return accuracy

# Entry to Task 4 until Task 9: Run the same prediction loop from Session 3, Wrap the prediction workflow in a function


def run_prediction_loop(dataset):
    """Run the same prediction loop from Session 3, but inside a function.

    This function keeps the beginner-friendly variables:
    correct, wrong, total, and y_pred_list.

    Args:
        dataset (list): A list of flower dictionaries.

    Returns:
        tuple: correct, wrong, total, y_pred_list, accuracy
    """
    # Task III from Session III: initialize metrics and predictions
    correct = 0      # Count of correct predictions
    wrong = 0        # Count of wrong predictions
    total = 0        # Total samples processed
    y_pred_list = []  # List of all predictions made

    print("\n=== Start session 4 Prediction Loop ===")

    # Task IV and Task V from Session III
    for sample in dataset:
        print("\nProcessing sample with id:", sample["id"])

        # Task 5: Predict the class using petal length
        y_pred = compute_threshold_prediction(sample)

        # Task 6: Convert species into the lesson’s binary label
        # For task 6, we purposely mention, they need to write do the y_true = derive_true_label(sample), but the function body of the derive_true_label is provided as it is.
        y_true = derive_true_label(sample)

        # Task 7: Update prediction counts and save results
        # For task 7, we will not provide the correct, wrong, total, y_pred_list = update_result_counts( correct, wrong, total, y_pred_list, y_pred, y_true ), even the function body is only a pass
        correct, wrong, total, y_pred_list = update_result_counts(
            correct, wrong, total, y_pred_list, y_pred, y_true
        )

        # Task 8: Display the result for each sample
        # The just need to uncomment the print statement below to see the result for each sample,
        print(
            f"id={sample['id']} | true={y_true} | pred={y_pred} | "
            f"petal_length={sample['petal_length']}"
        )
    # Task 9: Return the prediction loop results
    # student need to file <your_code>,wrong, <your_code>,<your_code>
    return correct, wrong, total, y_pred_list

# Task 10: Calculate accuracy


def calculate_accuracy(correct, total):
    """Calculate accuracy percentage.

    Args:
        correct (int): Number of correct predictions.
        total (int): Number of processed samples.

    Returns:
        float: Accuracy percentage.
    """
    if total > 0:
        accuracy = (correct / total) * 100
    else:
        accuracy = 0.0
    # Student need to provide the return value, maybe, we can ask them, what is missing in this function, go and check the function call in the main function, and see what they need to return here to make the main function work.
    return accuracy

# Task 12: Create and call the summary-report function


def print_summary(correct, wrong, total, y_pred_list, accuracy):
    """Print the final results after the loop is finished.

    Args:
        correct (int): Number of correct predictions.
        wrong (int): Number of wrong predictions.
        total (int): Number of processed samples.
        y_pred_list (list): List of predicted labels.
        accuracy (float): Accuracy percentage.
    """
    print("\n=== session 4 Summary ===")
    print("Correct:", correct)
    print("Wrong:", wrong)
    print("Total:", total)
    print("Accuracy (%):", round(accuracy, 2))
    print("All predictions:", y_pred_list)


def main():
    """Run the full beginner version of the program."""

    # Task 1 : Create a helper function to print status updates
    make_print_status("Build dataset")

    # Task 2: Create the flower dataset
    dataset = setup_application_list()

    # Task 3 : Show a status update before running prediction
    # Uncomment this line to see status messages
    make_print_status("Run prediction loop")

    # Entry to Task 4 until Task 9: Wrap the prediction workflow in a function
    # correct, <your_code>, total, y_pred_list = run_prediction_loop(dataset)
    correct, wrong, total, y_pred_list = run_prediction_loop(dataset)

    # Task 10: Calculate accuracy
    accuracy = calculate_accuracy(correct, total)

    # Task 11: Show a status update before printing the summary
    # we ask student to write "Print summary" as the status text, replacing the <your_code>.
    make_print_status("Print summary")

    # Task 12: Create and call the summary-report function
    print_summary(correct, wrong, total, y_pred_list, accuracy)
    print("\n=== End of session 4 ===")


if __name__ == "__main__":
    main()
]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader>TemplateGrader</grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>

    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="0.0000000">
        <testcode>
          <text>Check Session 4 Task 7 using the student's pasted code.</text>
        </testcode>
        <stdin>
          <text></text>
        </stdin>
        <expected>
          <text></text>
        </expected>
        <extra>
          <text></text>
        </extra>
        <display>
          <text>SHOW</text>
        </display>
      </testcase>
    </testcases>

    <tags>
      <tag>
        <text>session4</text>
      </tag>
      <tag>
        <text>submission</text>
      </tag>
      <tag>
        <text>ses4_task7_iris</text>
      </tag>
      <tag>
        <text>metrics-update</text>
      </tag>
    </tags>
  </question>
</quiz>
